Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = ""
COLLABORATORS = ""
```

---

For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

---
## Part A (2 points)

Write a function that finds the power $b^x$ for any positive integer x.


<!-- grade_id: exp_pos, points: 0, grade: false, solution: true -->
```python
def exp_pos(b,x):
    """Find the power of b^x. Assume x is a positive integer.   
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

For $n=2,x=3,$ your function should print 8. Check that it does:

```python
exp_pos(2,3)
```

<!-- grade_id: correct_exp_pos, points: 1, grade: true, solution: false -->
```python
"""Check that exp_pos returns the correct output for several inputs"""
assert exp_pos(3,2) == 9
assert exp_pos(3,4) == 81
assert exp_pos(1,4) == 1
assert exp_pos(4,1) == 4
assert exp_pos(0,4) == 0
```

<!-- grade_id: correct_-b_exp_pos, points: 1, grade: true, solution: false -->
```python

```

---

## Part B (3 points)

Write a function that find the power $b^x$ for any negative integer $x$. Make sure it handles the case where $b=0$ by raising a `ValueError`.

<!-- grade_id: exp_neg, points: 0, grade: false, solution: true -->
```python
def exp_neg(b,x):
    """Find the power of b^x. Assume x is a negative integer. 
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

For $n=2,x=-3,$ your function should print 0.125. Check that it does:

```python
exp_neg(2,-3)
```

<!-- grade_id: correct_exp_neg, points: 1, grade: true, solution: false -->
```python
"""Check that exp_neg returns the correct output for several inputs"""
assert exp_neg(3,-2) == 1/9
assert exp_neg(3,-4) == 1/81
assert exp_neg(1,-4) == 1
assert exp_neg(4,-1) == 1/4
```

<!-- grade_id: correct_-b_exp_neg, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: correct_exp_neg_error, points: 1, grade: true, solution: false -->
```python
"""Check that exp_neg raises an error for invalid inputs"""
try:
    exp_neg(0,-2)
except ValueError:
    pass
else:
    raise AssertionError("did not raise ValueError")
    
try:
    exp_neg(0,-1)
except ValueError:
    pass
else:
    raise AssertionError("did not raise ValueError")
```

---

## Part C (2 points)

Using `exp_pos` and `exp_neg` write a function to find the power $b^x$ for any integer $x$.

<!-- grade_id: sum_of_squares, points: 0, grade: false, solution: true -->
```python
def exp(b,x):
    """Find the power of $b^x$.  
    """
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: correct_exp, points: 1, grade: true, solution: false -->
```python
"""Check that exp returns the correct output for several inputs"""
assert exp(3,2) == 9
assert exp(3,4) == 81
assert exp(1,4) == 1
assert exp(4,1) == 4
assert exp(3,-2) == 1/9
assert exp(3,-4) == 1/81
assert exp(1,-4) == 1
assert exp(4,-1) == 1/4

"""Check that exp raises an error for invalid inputs"""
try:
    exp(0,-2)
except ValueError:
    pass
else:
    raise AssertionError("did not raise ValueError")
```

<!-- grade_id: correct_p0_exp, points: 1, grade: true, solution: false -->
```python

```