For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

---
## Part A (2 points)

Write a function that finds the power $b^x$ for any positive integer x.


<!-- grade_id: exp_pos, points: 0, grade: false, solution: true -->
```python
def exp_pos(b,x):
    """Find the power of b^x. Assume x is a positive integer.   
    """
    ### BEGIN SOLUTION
    p=b
    e=x
    while e>1:
        p*=b
        e-=1
    return p
    ### END SOLUTION
```

For $n=2,x=3,$ your function should print 8. Check that it does:

```python
exp_pos(2,3)
```

Output:
```
8
```

<!-- grade_id: correct_exp_pos, points: 1, grade: true, solution: false -->
```python
"""Check that exp_pos returns the correct output for several inputs"""
assert exp_pos(3,2) == 9
assert exp_pos(3,4) == 81
assert exp_pos(1,4) == 1
assert exp_pos(4,1) == 4c
assert exp_pos(0,4) == 0
```

Output:
```
File "<ipython-input-3-ad3e585ea929>", line 5
    assert exp_pos(4,1) == 4c
                            ^
SyntaxError: invalid syntax
```

<!-- grade_id: correct_-b_exp_pos, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert exp_pos(-2,3) ==-8
assert exp_pos(-1,2) ==1
assert exp_pos(-1,3) ==-1
### END HIDDEN TESTS
```

---

## Part B (3 points)

Write a function that find the power $b^x$ for any negative integer $x$. Make sure it handles the case where $b=0$ by raising a `ValueError`.

<!-- grade_id: exp_neg, points: 0, grade: false, solution: true -->
```python
def exp_neg(b,x):
    """Find the power of b^x. Assume x is a negative integer. 
    """
    ### BEGIN SOLUTION
    if b==0: 
        raise ValueError("Cannot raise base 0 to a negative exponent")
    p=1/b
    e=-x
    while e>1:
        p*=1/b
        e-=1
    return p
    ### END SOLUTION
```

For $n=2,x=-3,$ your function should print 0.125. Check that it does:

```python
exp_neg(2,-3)
```

<!-- grade_id: correct_exp_neg, points: 1, grade: true, solution: false -->
```python
"""Check that exp_neg returns the correct output for several inputs"""
assert exp_neg(3,-2) == 1/9
assert exp_neg(3,-4) == 1/81
assert exp_neg(1,-4) == 1
assert exp_neg(4,-1) == 1/4
```

<!-- grade_id: correct_-b_exp_neg, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert exp_neg(-2,-3) ==-1/8
assert exp_neg(-1,-2) ==1
assert exp_neg(-1,-3) ==-1
### END HIDDEN TESTS
```

<!-- grade_id: correct_exp_neg_error, points: 1, grade: true, solution: false -->
```python
"""Check that exp_neg raises an error for invalid inputs"""
try:
    exp_neg(0,-2)
except ValueError:
    pass
else:
    raise AssertionError("did not raise ValueError")
    
try:
    exp_neg(0,-1)
except ValueError:
    pass
else:
    raise AssertionError("did not raise ValueError")
```

---

## Part C (2 points)

Using `exp_pos` and `exp_neg` write a function to find the power $b^x$ for any integer $x$.

<!-- grade_id: sum_of_squares, points: 0, grade: false, solution: true -->
```python
def exp(b,x):
    """Find the power of $b^x$.  
    """
    ### BEGIN SOLUTION
    if x>=1: return exp_pos(b,x)
    if x<=-1: return exp_neg(b,x)
    else: return 1
    ### END SOLUTION
```

<!-- grade_id: correct_exp, points: 1, grade: true, solution: false -->
```python
"""Check that exp returns the correct output for several inputs"""
assert exp(3,2) == 9
assert exp(3,4) == 81
assert exp(1,4) == 1
assert exp(4,1) == 4
assert exp(3,-2) == 1/9
assert exp(3,-4) == 1/81
assert exp(1,-4) == 1
assert exp(4,-1) == 1/4

"""Check that exp raises an error for invalid inputs"""
try:
    exp(0,-2)
except ValueError:
    pass
else:
    raise AssertionError("did not raise ValueError")
```

<!-- grade_id: correct_p0_exp, points: 1, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
assert exp(2,0) == 1
assert exp(-2,0) == 1
### END HIDDEN TESTS
```